# Question 1

In this question, you'll write some coding that performs string manipulation. This is pretty much your warm-up.

Remember from the Assignment 0 instructional video: the first cell block is where you put your answer. There may already be some code there, in which case **don't remove it**; just add what is asked for.

The only code you should remove is the single line that starts with `raise`; delete this or comment it out.

All the cells under the first one are autograder cells; don't edit those! But you can run them (using the Run button at the top) to test your own code! If they don't spit out any error messages, then your code works!

### Part A

What's your favorite *positive* number? Create a variable named `favorite_number` and assign it to be your favorite *positive* number.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1a_test, points: 6, grade: true, solution: false -->
```python
print("My favorite number is: " + str(favorite_number))
assert favorite_number >= 0
```

### Part B

Print out a famous quote! In the code below, create two variables:
 - one named `famous_person`,
 - the other named `their_quote`,
 
and assign them to be the name of some famous person and a quote that they said, respectively. Remember: these are *strings*!

<!-- grade_id: q1b, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q1b_test, points: 6, grade: true, solution: false -->
```python
print("{}, at age {}, said:\n\n\"{}\"".format(famous_person, favorite_number, their_quote))

assert len(famous_person) > 0
assert len(their_quote) > 0
```

### Part C

You're working late on a homework assignment and have copied a few lines from a Wikipedia article. In your tired stupor, your copy/paste skills leave something to be desired. Rather than try to force your mouse hand to stop shaking, you figure it's easier to write a small Python program to strip out errant whitespace from your copying-pasting.

**Leaving the code below as-is** (other than the line beginning with `raise`, of course), **add more code** that strips away whitespace in the variables `line1`, `line2`, and `line3`.

In other words, reassign each of the three variables (keeping their names the same, and their respective content otherwise unaltered) so that the leading and trailing whitespace on both ends is removed.

<!-- grade_id: q1c, points: 0, grade: false, solution: true -->
```python
line1 = 'Python supports multiple programming paradigms, including object-oriented, imperative\n'
line2 = '  and functional programming or procedural styles. It features a dynamic type\n'
line3 = ' system and automatic memory management and has a large and comprehensive standard library.\n '

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: q1c_test, points: 6, grade: true, solution: false -->
```python
assert line1[-1] != '\n'
assert line2[0] != line2[1] != ' '
assert line2[-1] != '\n'
assert line3[0] != ' '
assert line3[-1] != '\n'
```

### Part D

You discover that there are numbers in the text you'd like to be able to parse out for some math down the road. After stripping out the trailing whitespace, convert the numbers to their proper numerical form.

Assign them to the variables `num1`, `num2`, and `num3` respectively.

<!-- grade_id: q1d, points: 0, grade: false, solution: true -->
```python
line1 = ' 495.59863 \n'
line2 = '\t134  '
line3 = '\n\t -5.4 \t'

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: q1d_test, points: 6, grade: true, solution: false -->
```python
assert num1 > 495 and num1 < 496
assert num2 == 134
assert num3 > -6 and num3 < -5
```

### Part E

Take the `number` below, find its square root, convert it to a string, and then assign it to a variable named `str_version`.

You **must use the correct arithmetic operator for the square root, as well as the correct casting function for the string conversion**.

**No imports allowed!**

<!-- grade_id: q1e, points: 0, grade: false, solution: true -->
```python
number = 3.14159265359

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: q1e_test, points: 6, grade: true, solution: false -->
```python
assert str_version == "1.7724538509055743"
```