# Question 2

This question will walk you through the process of how a coding assignment in JupyterHub will work. In the following parts, you'll edit the sections of code needed to pass the autograder.

For each part below, you'll notice there's a cell with one line of Python, starting with the word `raise`. **Delete this line**, and in its place, put the code that answers the prompt.

Below that cell, you'll find another cell with multiple lines of Python, starting with `try:`. **DO NOT EDIT THIS CELL.** This is the autograder. For future reference, JupyterHub knows what the autograder code is before you download the assignment, so even if you edit the autograder so your code passes its tests, the edits will be removed when you submit the assignment. So you can't cheat that way :)

**HINT:** Every time you fill in code to answer a question, you can immediately test that code by selecting the cell and pressing the "Run Cell" button in the menu (looks like a Play icon). After you run the cell with your code, repeat this process with the autograder cell(s) below it. Provided there aren't any errors, your code worked correctly! If there are errors, go back and edit your code and repeat this process.

### Part A

Create a variable `x`, and *assign* it to have the value `3`.

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q2a_test, points: 10, grade: true, solution: false -->
```python
assert x == 3
```

### Part B

Reset (or *re-assign*) the variable `x` to have the value `3.14159`.

<!-- grade_id: q2b, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q2b_test, points: 10, grade: true, solution: false -->
```python
assert x == 3.14159
```

### Part C

Create a new variable `y`, and assign it to have the value `10`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q2c_test, points: 10, grade: true, solution: false -->
```python
assert y == 10
```

### Part D

Create a new variable `z`, and assign it to have the product of the variables `x` and `y`.

<!-- grade_id: q2d, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q2d_test, points: 10, grade: true, solution: false -->
```python
assert z == 31.4159
```

### Part E

Create a new variable `z_squared`, and assign it to have the value of `z` raised the power 2.

<!-- grade_id: q2e, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: q2e_test, points: 10, grade: true, solution: false -->
```python
assert z_squared == 986.95877281
```

### Part F

With that done, all that's left is to submit your code. First, make sure you save your work (go to "File" and choose "Save and Checkpoint".

Go back to the "Assignments" tab ("File" -> "Close and Halt", then click the "Assignments" tab). Notice that, in the "Downloaded assignments" tab, there's a little arrow next to "A0". Click it, and a drop-down will appear with a handy "Validate" button on the far right. This button is awesome--it runs all the autograder tests before you even submit! So you automatically know what [autograded] questions work, and which don't. Of course, there can (and will!) still be manually-graded questions, but it should give you a good idea how you're doing.

Go ahead and click validate. Provided you followed the directions, you should see this popup:

![successkid](http://cs.uga.edu/~squinn/courses/su16/csci1360e/assets/successkid.png)

All tests passed, well done! All that's left is to send the assignment back to me so I can officially record your grade. Hit the blue "Submit" button on the right. Once it goes through, notice the timestamp that shows up:

![timestamp](http://cs.uga.edu/~squinn/courses/su16/csci1360e/assets/timestamp.png)

This is the server's officially-recorded time of completion for your assignment. You can submit as many times as you want! Only the last submission will ever be graded.

**This step is crucial. If you don't submit, I will have no way of knowing that you completed your assignment.**

### In the box below, tell me whether or not you were able to submit successfully (and get the timestamp confirmation).

<!-- grade_id: q2f, points: 5, grade: true, solution: true -->
